# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
file = pd.read_csv('assets/NISPUF17.csv')
data = pd.DataFrame(file)
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE

    n = data['EDUC1'].value_counts()/data['EDUC1'].shape[0]
    # print(n)
    name_list = ['less than high school', 'high school', 'more than high school but not college', 'college']
    dictionary = {}
    i = 0
    for label in name_list:
        dictionary[label] = n[i+1] 
        i += 1
    
    return dictionary


print(proportion_of_education())
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

```python
import pandas as pd
file = pd.read_csv('assets/NISPUF17.csv')
data = pd.DataFrame(file)
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
     # import pandas as pd
    # mb = pd.read_csv("assets/NISPUF17.csv")
#     n = data['EDUC1'].value_counts()/data['EDUC1'].shape[0]
#     print(n)
#     li = ['less than high school', 'high school', 'more than high school but not college', 'college']
#     di = {}
#     i = 0
#     for l in li:
#         di[l] = n[i+1] 
#         i += 1
    
#     return di


# print(proportion_of_education())
    not_high_school= high_school= not_college= college = 0
    for val_1 in data['EDUC1']:
        val = int(val_1)
        if val==1:
            not_high_school+=1
        elif val==2:
            high_school+=1
        elif val==3:
            not_college+=1
        else:
            college+=1
            
    not_high_school = not_high_school/data.shape[0]
    high_school = high_school/data.shape[0]
    not_college = not_college/data.shape[0]
    college = college/data.shape[0]
    
    proportion = {"less than high school":not_high_school,
    "high school":high_school,
    "more than high school but not college":not_college,
    "college":college} 
    return proportion
#     raise NotImplementedError()

print(proportion_of_education())

```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    
    file = pd.read_csv('assets/NISPUF17.csv')
    data = pd.DataFrame(file)
    
    # CBF_01  = column with child fed breast milk yes/no
    # P_NUMFLU  = column with total number of seasonal infulenza doses
              
    data = data[['CBF_01','P_NUMFLU']]
    
    milk_fed_vaccine = milk_not_fed_vaccine = milk_fed_child = milk_not_fed_child = 0
   
    # eliminating rows having NaN values in column P_NUMFLU and keeping the remaining rows
    data = data[pd.notnull(data['P_NUMFLU'])]
    
    for row in zip(data['CBF_01'], data['P_NUMFLU']):
        
        if int(row[0])==1: # child fed breast milk 
            milk_fed_vaccine = milk_fed_vaccine + int(row[1])
            milk_fed_child+=1
        
        if int(row[0])==2: # child was not fed breast milk 
            milk_not_fed_vaccine = milk_not_fed_vaccine + int(row[1])
            milk_not_fed_child+=1
            
    return tuple((milk_fed_vaccine/milk_fed_child , milk_not_fed_vaccine/milk_not_fed_child))
    raise NotImplementedError()
 
    """
    milk_fed_vaccine = 24986
    ,milk_not_fed_vaccine = 3188
    ,milk_fed_child = 13291
    ,milk_not_fed_child = 1997
    
    average = (1.8799187420058687, 1.5963945918878317)
    
    """
print(average_influenza_doses())

```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    file = pd.read_csv('assets/NISPUF17.csv')
    data = pd.DataFrame(file)
    
    # HAD_CPOX  = did child ever had chicken pox
    # P_NUMVRC = total number of varicella doses
    
    data = data[['HAD_CPOX','P_NUMVRC','SEX']]
    data = data[pd.notnull(data['P_NUMVRC'])]
    
    male_pox =  0 
    male_not_pox = 0
    female_pox = 0
    female_not_pox = 0
    c=0
    for row in zip(data['P_NUMVRC'], data['HAD_CPOX'], data['SEX']):
        
        if int(row[0])==1  and row[1]==1: # child had chicken pox
            if int(row[2]) == 1:
                male_pox+=1
            if int(row[2]) == 2:
                female_pox+=1
        
        if int(row[0])==1 and row[1]==2: # child did not had chicken pox
            if int(row[2]) == 1:
                male_not_pox+=1
            if int(row[2]) == 2:
                female_not_pox+=1
    
    # print(male_pox , male_not_pox , female_pox , female_not_pox ) 
    
    return {"male": male_pox/male_not_pox,"female":female_pox/female_not_pox}
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                   "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
import pandas as pd
data = pd.read_csv('assets/NISPUF17.csv')
v1 =  data[(data['P_NUMVRC'] >=0) & (data['HAD_CPOX'] <= 2)]

print(v1)



```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
0               1   128521     12852     2         NaN  235.916956     1031   
1               2    10741      1074     2         NaN  957.353840     1068   
2               3   220011     22001     2         NaN  189.611299     1050   
3               4    86131      8613     1  675.430817  333.447418     1040   
4               5   227141     22714     1  482.617748  278.768063     1008   
...           ...      ...       ...   ...         ...         ...      ...   
28460       28461    19141      1914     2         NaN   21.188088     2062   
28461       28462   121401     12140     2         NaN   26.503010     2004   
28462       28463   128831     12883     2         NaN   27.520652     2004   
28463       28464    82921      8292     2         NaN   22.390587     2062   
28464       28465   244621     24462     2         NaN   27.520652     2004   

       YEAR  AGECPOXR  HAD_CPOX  ...  XVRCTY2  XVRCTY3  XV
```

```python

```